# This is a list of rule files to consider for testing. Globs are supported.
rule_files:
  - '../rules/generic/up.rules.yml'

# All the tests are listed here.
tests:
  - interval: 1m
    input_series:
      - series: up{instance="https://example.com", job="blackbox"}
        values: 0+0x30

    alert_rule_test:
      - eval_time: 30m
        alertname: TargetDown
        exp_alerts:
          - exp_labels:
              severity: critical
              paging: true
              priority: P2
              job: blackbox
              instance: https://example.com
            exp_annotations:
              summary: 'Target is down: blackbox on https://example.com'
              description: "blackbox on https://example.com is not reachable"

  - interval: 1m
    input_series:
      - series: >-
          probe_success{instance="https://example.com",
          job="blackbox", module="http"}
        values: 0+0x5

    alert_rule_test:
      - eval_time: 5m
        alertname: BlackboxProbeFailed
        exp_alerts:
          - exp_labels:
              severity: critical
              paging: true
              priority: P1
              job: blackbox
              module: http
              instance: https://example.com
            exp_annotations:
              summary: 'Blackbox check failed http on https://example.com'
              description: "http on https://example.com is not reachable"


  - alert_rule_test:
      - eval_time: 5m
        alertname: NoTargetFound
        exp_alerts:
          - exp_labels:
              severity: critical
              paging: true
              priority: P1
            exp_annotations:
              summary: No scrape target could be found
              description: >-
                All configured scrape targets are down, please have a look
                if prometheus is working properly.
